Before you submit this notebook, make sure everything runs as expected in the local test cases. 
Please, paste the solution to the designed cell and do not change anything else.

Also, please, leave your first and last names below

```python
FirstName = "Arseny"
LastName = "Staroverov"
```

---

```python
from sklearn.base import BaseEstimator, ClassifierMixin
from sklearn.metrics import accuracy_score

import numpy as np
import torch
import torch.nn as nn
import torch.optim as optim

from scipy.spatial.distance import cdist # You can use this for rbf kernel

import unittest
import sys
import io
```

<!-- grade_id: cell-7bfbe050819ecfac, points: 0, grade: false, solution: true -->
```python
def rbf(x_1, x_2, sigma=1.):
    '''Computes rbf kernel for batches of objects

    Args:
        x_1: torch.tensor shaped `(#samples_1, #features)` of type torch.float32
        x_2: torch.tensor shaped `(#samples_1, #features)` of type torch.float32
    Returns:
        kernel function values for all pairs of samples from x_1 and x_2
        torch.tensor of type torch.float32 shaped `(#samples_1, #samples_2)`
    '''
    distances = torch.exp(-torch.cdist(x_1, x_2) ** 2 / (2 * sigma ** 2))
    return torch.Tensor(distances).type(torch.float32)
```

### Test 0: Initialization (0.01 points)

<!-- grade_id: cell-68578dd758ce6174, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
from sklearn.datasets import make_circles
X, y = make_circles(150, factor=.1, noise=.1, random_state=42)


```

### Test 1: SVM accuracy (0.69 points)

<!-- grade_id: cell-1759990cd7cea5ce, points: 0.69, grade: true, solution: false -->
```python
X_test, y_test = X[100:], y[100:]
X, y = X[:100], y[:100]
y[y==0] = -1 # for convenience with formulas
y_test[y_test==0] = -1
clf = SVM(epochs=150, lr=0.1, batch_size=20, verbose=False, kernel_function=rbf)
clf.fit(X, y)
pred = clf.predict(X_test)
assert accuracy_score(y_test, pred) > 0.95
```

### Test 2: Kernel (0.3 points)

<!-- grade_id: cell-a92ec82e4098bdd2, points: 0.3, grade: true, solution: false -->
```python
assert (np.allclose(rbf(X_kernel, X_kernel).detach().cpu().numpy(), np.array([[1.0000e+00, 2.2897e-11],[2.2897e-11, 1.0000e+00]])))
```